# Relación de ejercicios 1

### Lee detenidamente las siguientes instrucciones y borra `raise NotImplementedError` en todas las celdas

En cada trozo de código debes responder a la pregunta formulada, asegurándote de que el resultado queda guardado en la(s) variable(s) indicadas, que por defecto vienen inicializadas a `None`. Puedes usar variables intermedias siempre que el resultado final del cálculo quede guardado exactamente en la variable que indica el enunciado (debes reemplazar `None` el código necesario, pero nunca cambiar el nombre de las variables ya creadas). 

Después de cada ejercicio verás varias líneas de código ya hechas. Ejecútalas todas (no modifiques su código) y te dirán si tu solución es correcta o no, sin darte pistas de cómo se resuelve (el código de los tests no guarda relación con el código de la solución). Si la solución es correcta, no se mostrará nada, pero si es incorrecta, verás un error indicando cuál es el test que tu solución no pasa. Además de esas pruebas, se realizarán algunas más (ocultas) a la hora de puntuar el ejercicio, pero si tu código pasa con éxito las líneas que ves, puedes estar bastante seguro de que tu solución es correcta. Asegúrate de que todas las celdas de auto-evaluación indican que tu solución es correcta antes de subir el ejercicio a la plataforma. En caso contrario, no dudes en preguntar las dudas a través del foro -siempre explicando la duda en palabras, sin pegar código fuente-.

Una vez finalizada la actividad, guarda tu fichero, ciérralo con File -> Close and Halt, y después vuélvelo a abrir y ejecútalo completo pinchando en Cell -> Run All Cells, y asegúrate de que no se lanza ningún error. De esta manera comprobarás que no has olvidado nada y que es posible ejecutarlo completo desde 0 y sin errores. No se corregirá ningún fichero que tenga errores de sintaxis y no se pueda, como mínimo, ejecutar completamente. No pasa nada si alguna de las comprobaciones lanza errores por ser incorrecta, pero el código de cada ejercicio no puede tener errores de sintaxis y debe al menos poderse ejecutar.

RECUERDA SUBIR CADA UNO DE LOS FICHEROS .ipynb TAL CUAL (sueltos), SIN COMPRIMIR Y SIN CAMBIARLES EL NOMBRE. Los ficheros subidos deben tener **exactamente el mismo nombre de fichero que tenían cuando los recibiste**. No subas ningún PDF ni ningún fichero ZIP ni nada similar. La plataforma ya los separa automáticamente en carpetas que traen el nombre y apellidos del alumno, por lo que NO es necesario que lo pongas en ninguna parte.

**Ejercicio 1 (2 puntos)**

* Crear en la variable `lista` una **lista** de enteros empezando en 3, no mayores que 180 en pasos de 4 elementos
* Crear en la variable `arr` un ndarray con la función `np.array()` y a partir de él, crear en `mascara` una máscara booleana de los elementos que son múltiplos de 3. **No está permitido utilizar bucles**. No olvides los `import` necesarios
* Calcula en la variable `n_multiplos` el número total de elementos múltiplos de 3 que has obtenido. *No está permitido asignar el valor a mano*.
* Utilizando la máscara, reemplazar en la propia variable `arr` los elementos múltiplos de 3 por el doble de su valor. **No está permitido utilizar bucles**.
* Crear en la variable `matriz` un array (una matriz) de 3 filas que deben ser el vector `arr` tras el reemplazamiento, el doble de dicho vector, y la mitad del mismo, respectivamente. PISTA: utiliza la misma función que has empleado para crear el array, pasándole una lista formada por los 3 vectores indicados, tendrás la matriz. Comprueba sus dimensiones con `shape` después.

<!-- grade_id: ej1_solucion, points: 0, grade: false, solution: true -->
```python
import numpy as np


#a)
lista = list(range(3,181,4))
#b)
arr = np.array(lista)
mascara = arr%3 == 0
# seleccionar las posiciones que sean True para visualizar los valores

num_mascara =  arr[mascara]
n_multiplos =  len(num_mascara)

np.multiply.at(arr, (mascara), 2)

# c) 
fila2 = arr*2
fila3 = arr/2

matriz = np.array([arr, fila2,fila3])
matriz.sum()
len(lista)

```

Output:
```
45
```

<!-- grade_id: ej1_test, points: 2, grade: true, solution: false -->
```python
import numpy as np
assert(isinstance(lista, list))
assert(len(lista) == 45)
assert(min(lista) == 3 and max(lista) == 179 and sum(lista) == 4095)   # el contenido de la lista es incorrecto
assert(isinstance(arr, np.ndarray) and arr.size == 45)
assert(isinstance(mascara, np.ndarray) and mascara.size == 45 and mascara.dtype.name == "bool")  # array booleano
assert(n_multiplos == 15)
assert(isinstance(matriz, np.ndarray) and matriz.shape == (3, 45))
assert(arr.sum() == 5400)             # suma de arr tras haber reemplazado los múltiplos de 3 por el doble de su valor
assert(matriz.sum() == 18900)
```

**Ejercicio 2 (2 puntos)**

* **Sin utilizar bucles** ni map, sólo operaciones con vectores de numpy, crear una función llamada `modulo` que reciba un solo argumento `v`, que será un vector de numpy, y calcule el módulo de dicho vector, esto es, su magnitud, tal como indica la siguiente fórmula:

Dado un vector n-dimensional $v=(v_1,...,v_n)$, se define el módulo de v como 

$$|v| = \sqrt{v_1^2 + ... + v_n^2}$$

La función raíz cuadrada puede aplicarse a un vector si usas la versión que hay en numpy, `np.sqrt` y está vectorizada, es decir, puede recibir un vector y calculará otro vector con la raíz de cada elemento. Lo mismo ocurre con el operador potencia `**` que está vectorizado se puede aplicar elemento a elemento. Recuerda que los vectores de numpy pueden devolver la suma de sus elementos.

* Utilizando la función `np.apply_along_axis`, invocar a la función `modulo` sobre cada fila de la matriz `matriz` definida en el ejercicio anterior, almacenando el resultado en la variable `modulos_filas` que al final contendrá un vector de 3 elementos. PISTA: `np.apply_along_axis(funcion, axis, matriz)` sirve para aplicar una función definida por el usuario, que *debe estar pensada para recibir como argumento un vector*, a cada una de las dimensiones de un array. Debemos fijar `axis=1` para aplicarla a cada fila, lo cual resultará en un vector de 3 elementos formado por los resultados de invocar 3 veces a la función (sobre cada una de las 3 filas de la matriz).

<!-- grade_id: ej2_solucion, points: 0, grade: false, solution: true -->
```python

import math

v = np.array([2,3])
#modulo = modulo(v)

def modulo(v):
    return np.sqrt(sum(pow(v, 2)))

print('Vector:', v)
print('Magnitude of the Vector:', modulo(v))
modulo(v)

modulos_filas = np.apply_along_axis(modulo, axis = 1, arr = matriz)
modulos_filas
```

Output:
```
Vector: [2 3]
Magnitude of the Vector: 3.605551275463989
array([ 977.57352665, 1955.14705329,  488.78676332])
```

<!-- grade_id: ej2_test, points: 2, grade: true, solution: false -->
```python
assert(isinstance(modulos_filas, np.ndarray) and modulos_filas.size == 3)
assert(round(modulos_filas.mean(), 2) == 1140.5)
```

**Ejercicio 3 (2 puntos)**

Crear una función `quita_primer_valor` que reciba un argumento llamado `d` que será un diccionario donde todas las claves serán de tipo string y el valor asociado a cada clave será de tipo lista. La función debe eliminar el primer elemento de cada lista y debe devolver un nuevo diccionario con el resultado de dicha modificación (por tanto, el diccionario que le hemos pasado como argumento no debe ser modificado). Asumir que las listas siempre tendrán al menos un elemento.
* PISTA: utiliza **listas por comprensión** para crear una nueva lista de pares a partir de la lista formada por los pares del diccionario original. La nueva lista de pares creada debe tener en cada tupla su segundo elemento modificado, como el resultado de tomar un segmento (de la posición 1 hasta el final) de la lista original que ya existía asociada a esa clave. Finalmente, crea un nuevo diccionario con la functión `dict` a partir de dicha lista de pares, y devuélvelo.
* Puedes probar la función sobre el diccionario `midict = {"edad": [22, 45, 32], "nombre": ["Juan", "Paula", "Elena"]}`.

<!-- grade_id: ej3_solucion, points: 0, grade: false, solution: true -->
```python
midict = {"edad": [22, 45, 32], "nombre": ["Juan", "Paula", "Elena"]}
lista_ini = list(midict.values())
lista_ini
def quita_primer_valor(midict):
    palabras_less =[]
    for p in lista_ini:
        palabras_less.append(p[1:])
    diccionario_final = {"edad": palabras_less[0],"nombre": palabras_less[1]}
    return diccionario_final

diccionario_final =quita_primer_valor(midict)

diccionario_final
```

Output:
```
{'edad': [45, 32], 'nombre': ['Paula', 'Elena']}
```

<!-- grade_id: ej3_test, points: 2, grade: true, solution: false -->
```python
midict = {"edad": [22, 45, 32], "nombre": ["Juan", "Paula", "Elena"]}
res = quita_primer_valor(midict)
assert(isinstance(res, dict) and len(res.items()) == 2)
assert(res["edad"] == [45, 32] and res["nombre"] == ["Paula", "Elena"])
```

**Ejercicio 4 (2 puntos)**

Implementa una función `quita_menores` que reciba dos argumentos: el primero llamado `personas` debe ser una lista de diccionarios, cada uno con datos personales de una persona, entre los cuales se encuentra el campo `"edad"`. El segundo, llamado `edad_minima`, debe ser un entero y debe tener valor por defecto 18 si no se especifica dicho argumento. La función debe devolver una nueva lista en la que se han descartado aquellas personas cuyo campo `edad` sea menor que `edad_limite`. Puedes probar tu función con la lista `personas = [{"nombre": "Pepe", "edad": 15}, {"nombre": "Juan", "edad": 59}, {"nombre": "Marisa", "edad": 42}, {"nombre": "Eva"}]`.
* PISTA: utiliza **listas por comprensión con un if dentro**, para recorrer la lista de personas y construir otra lista en la que sólo estará cada persona si su edad es mayor o igual que la edad mínima indicada como argumento.
* PISTA: tener en cuenta que alguno de los diccionarios puede no contener el campo "edad" en cuyo caso dicha persona no será incluida en la lista resultante, pero *el código no debe lanzar ningún error* en dicho caso. Recuerda la función `get()` de los diccionarios.

<!-- grade_id: ej4_solucion, points: 0, grade: false, solution: true -->
```python
personas = [{"nombre": "Pepe", "edad": 15},
            {"nombre": "Juan", "edad": 59},
            {"nombre": "Marisa", "edad": 42}, 
            {"nombre": "Eva"}] 

def quita_menores(personas, edad_minima=18):
    return[item for item in personas if item.get("edad",0) >=edad_minima]
quita_menores(personas, 18)
```

Output:
```
[{'nombre': 'Juan', 'edad': 59}, {'nombre': 'Marisa', 'edad': 42}]
```

<!-- grade_id: ej4_test, points: 2, grade: true, solution: false -->
```python
personas = [{"nombre": "Pepe", "edad": 15}, {"nombre": "Juan", "edad": 59}, {"nombre": "Marisa", "edad": 42}, {"nombre": "Eva"}]
res_18 = quita_menores(personas)
assert(isinstance(res_18, list) and len(res_18) == 2)
assert(isinstance(res_18[0], dict) and isinstance(res_18[1], dict))
assert(res_18[0]["nombre"] == "Juan" and res_18[1]["nombre"] == "Marisa")

res_50 = quita_menores(personas, 50)
assert(isinstance(res_50, list) and len(res_50) == 1)
assert(isinstance(res_50[0], dict))
assert(res_50[0]["nombre"] == "Juan")
```

**Ejercicio 5 (2 puntos)**

Implementa una función llamada `reemplaza_max` que reciba dos argumentos: 
* Un argumento `m` que será una matriz numérica bidimensional de numpy (array bidimensional)
* Un segundo argumento `repl` que será un número real

La función debe devolver **una nueva matriz** en la que se ha reemplazado en cada fila de la matriz el máximo de esa fila por el valor del argumento `repl`. Para ello, aplica de nuevo la función `np.apply_along_axis(funcion, axis, matriz)` sobre una copia de la matriz `m`, con argumento `axis=1` (que significa "aplicar a cada fila") donde la función del primer argumento debe ser **una función lambda** de un solo argumento `v` (vector numérico de numpy), que devuelva un nuevo vector resultado de reemplazar en el vector fila original el máximo de ese vector por `repl`. Dentro de dicha función lambda puedes hacer referencia a la variable `repl` de la función externa sin ningún problema.

Puedes probar tu función sobre la matriz `matriz` definida en el Ejercicio 1, con `repl=-10`.

* PISTA: dentro de la función, crea en primer lugar una variable llamada `copia_m` que contenga una copia de `m` (método `m.copy()`), después invoca a `np_apply_along_axis` y finalmente devuelve `copia_m` como resultado de la función `reemplaza_maximos`.
* PISTA: Python no permite utilizar el operador de asignación dentro de una función lambda, pero puedes usar la función `np.put(v, indice, elemento)` que se comporta exactamente como si fuera `v[indice] = elemento`. Esto modifica directamente el vector pasado como argumento a la función lambda (es decir, cada fila de la matriz), lo cual no es ningún problema si invocas a la función `np_apply_along_axis` sobre una **copia** de la matriz *m* (y no sobre la matriz original que ha recibido `reemplaza_maximos`). Para saber el índice donde se encuentra el máximo del vector puedes usar `argmax()` aplicada al array.

```python
def reemplaza_max(m, repl):
    copia_m = m.copy()
    np.apply_along_axis(lambda v: np.put(v, v.argmax(), repl), 1, copia_m)
    return copia_m

reemplaza_max(matriz, -10)

```

Output:
```
array([[  6. ,   7. ,  11. ,  30. ,  19. ,  23. ,  54. ,  31. ,  35. ,
         78. ,  43. ,  47. , 102. ,  55. ,  59. , 126. ,  67. ,  71. ,
        150. ,  79. ,  83. , 174. ,  91. ,  95. , 198. , 103. , 107. ,
        222. , 115. , 119. , 246. , 127. , 131. , 270. , 139. , 143. ,
        294. , 151. , 155. , 318. , 163. , 167. , -10. , 175. , 179. ],
       [ 12. ,  14. ,  22. ,  60. ,  38. ,  46. , 108. ,  62. ,  70. ,
        156. ,  86. ,  94. , 204. , 110. , 118. , 252. , 134. , 142. ,
        300. , 158. , 166. , 348. , 182. , 190. , 396. , 206. , 214. ,
        444. , 230. , 238. , 492. , 254. , 262. , 540. , 278. , 286. ,
        588. , 302. , 310. , 636. , 326. , 334. , -10. , 350. , 358. ],
       [  3. ,   3.5,   5.5,  15. ,   9.5,  11.5,  27. ,  15.5,  17.5,
         39. ,  21.5,  23.5,  51. ,  27.5,  29.5,  63. ,  33.5,  35.5,
         75. ,  39.5,  41.5,  87. ,  45.5,  47.5,  99. ,  51.5,  53.5,
        111. ,  57.5,  59.5, 123. ,  63.5,  65.5, 135. ,  69.5,  71.5,
    
```

```python

def reemplaza_max(m, repl):
    return np.apply_along_axis(lambda v: np.where(v == v.max(), repl, v), 1, m.copy())

reemplaza_max(matriz, -10)
```

Output:
```
array([[  6. ,   7. ,  11. ,  30. ,  19. ,  23. ,  54. ,  31. ,  35. ,
         78. ,  43. ,  47. , 102. ,  55. ,  59. , 126. ,  67. ,  71. ,
        150. ,  79. ,  83. , 174. ,  91. ,  95. , 198. , 103. , 107. ,
        222. , 115. , 119. , 246. , 127. , 131. , 270. , 139. , 143. ,
        294. , 151. , 155. , 318. , 163. , 167. , -10. , 175. , 179. ],
       [ 12. ,  14. ,  22. ,  60. ,  38. ,  46. , 108. ,  62. ,  70. ,
        156. ,  86. ,  94. , 204. , 110. , 118. , 252. , 134. , 142. ,
        300. , 158. , 166. , 348. , 182. , 190. , 396. , 206. , 214. ,
        444. , 230. , 238. , 492. , 254. , 262. , 540. , 278. , 286. ,
        588. , 302. , 310. , 636. , 326. , 334. , -10. , 350. , 358. ],
       [  3. ,   3.5,   5.5,  15. ,   9.5,  11.5,  27. ,  15.5,  17.5,
         39. ,  21.5,  23.5,  51. ,  27.5,  29.5,  63. ,  33.5,  35.5,
         75. ,  39.5,  41.5,  87. ,  45.5,  47.5,  99. ,  51.5,  53.5,
        111. ,  57.5,  59.5, 123. ,  63.5,  65.5, 135. ,  69.5,  71.5,
    
```

<!-- grade_id: ej5_test, points: 2, grade: true, solution: false -->
```python
import inspect
import numpy as np
reemplaza_maximos_source = inspect.getsource(reemplaza_maximos)
m_antes = matriz.copy()
assert("lambda" in reemplaza_maximos_source and       # obligatorio utilizar una función lambda 
       "put(" in reemplaza_maximos_source and         # obligatorio utilizar la función np.put
       "copy()" in reemplaza_maximos_source)          # obligatorio utilizar copy() para no modificar el argumento
nueva_m = reemplaza_maximos(matriz, -10)
assert(nueva_m is not None)
assert((matriz == m_antes).all())                     # la matriz original pasada como argumento no puede ser modificada
assert((nueva_m[:, 42] == [-10, -10, -10]).all())
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In [26], line 7
      3 reemplaza_maximos_source = inspect.getsource(reemplaza_maximos)
      4 m_antes = matriz.copy()
      5 assert("lambda" in reemplaza_maximos_source and       # obligatorio utilizar una función lambda 
      6        "put(" in reemplaza_maximos_source and         # obligatorio utilizar la función np.put
----> 7        "copy()" in reemplaza_maximos_source)          # obligatorio utilizar copy() para no modificar el argumento
      8 nueva_m = reemplaza_maximos(matriz, -10)
      9 assert(nueva_m is not None)

AssertionError:
```